# Part III: Explore

![alt text](data/inu_neko_logo_small.png "Inu + Neko")

Hello! We are Inu + Neku and we are a Dog & Cat services and supplies store located in New York City. We just started our e-commerce business and need your help analyzing our data!

## Description

We need to make sure the data is clean before starting your analysis. As a reminder, we should check for:

- Duplicate records
- Consistent formatting
- Missing values
- Obviously wrong values

```python
import pandas as pd
import numpy as np
import matplotlib.pyplot as plt
```

```python
df_cleaned = pd.read_csv('data/inu_neko_orderline_clean.csv')
df_cleaned
```

Output:
```
trans_id      prod_upc  cust_id             trans_timestamp  \
0      10300097  719638485153  1001019  2021-01-01 07:35:21.439873   
1      10300093   73201504044  1001015  2021-01-01 09:33:37.499660   
2      10300093  719638485153  1001015  2021-01-01 09:33:37.499660   
3      10300093  441530839394  1001015  2021-01-01 09:33:37.499660   
4      10300093  733426809698  1001015  2021-01-01 09:33:37.499660   
...         ...           ...      ...                         ...   
38218  10327860  287663658863  1022098  2021-06-30 15:37:12.821020   
38219  10327960  140160459467  1022157  2021-06-30 15:45:09.872732   
38220  10328009  425361189561  1022189  2021-06-30 15:57:44.295104   
38221  10328089  733426809698  1022236  2021-06-30 15:59:29.801593   
38222  10328109  717036112695  1011924  2021-06-30 17:30:52.205912   

       trans_year  trans_month  trans_day  trans_hour  trans_quantity  \
0            2021            1          1           1               1   
1            2021   
```

```python
df_cleaned.info()
```

Output:
```
<class 'pandas.core.frame.DataFrame'>
RangeIndex: 38223 entries, 0 to 38222
Data columns (total 16 columns):
 #   Column            Non-Null Count  Dtype  
---  ------            --------------  -----  
 0   trans_id          38223 non-null  int64  
 1   prod_upc          38223 non-null  int64  
 2   cust_id           38223 non-null  int64  
 3   trans_timestamp   38223 non-null  object 
 4   trans_year        38223 non-null  int64  
 5   trans_month       38223 non-null  int64  
 6   trans_day         38223 non-null  int64  
 7   trans_hour        38223 non-null  int64  
 8   trans_quantity    38223 non-null  int64  
 9   cust_age          38223 non-null  int64  
 10  cust_state        38223 non-null  object 
 11  prod_price        38223 non-null  float64
 12  prod_title        38223 non-null  object 
 13  prod_category     38223 non-null  object 
 14  prod_animal_type  38223 non-null  object 
 15  total_sales       38223 non-null  float64
dtypes: float64(2), int64(9), object(5)
memor
```

#### Question 1: Number of Orders

How many transactions are there?

<!-- grade_id: cell-976c898e7f2b44c4, points: 0, grade: false, solution: true -->
```python
# your code here

num_trans = len(df_cleaned)
```

<!-- grade_id: cell-25ca64f53b142f34, points: 4, grade: true, solution: false -->
```python
# Q1 Test Cases
assert 39000 > num_trans > 25000
```

#### Question 2: Alpha and Omega I
What was the month and day of the first sale? Store as a tuple in that order and assign the tuple to the variable `first_date`.

<!-- grade_id: cell-bd2b89a8035a9ab8, points: 0, grade: false, solution: true -->
```python
# your code here
mth = df_cleaned.loc[0]["trans_month"]
day = df_cleaned.loc[0]["trans_day"]
first_date = (mth,day)
first_date
```

Output:
```
(1, 1)
```

<!-- grade_id: cell-8a2534037c10eb71, points: 4, grade: true, solution: false -->
```python
# Q2 Test Cases
assert len(first_date) == 2
assert str(first_date[0]).isnumeric() 
assert str(first_date[1]).isnumeric() 
```

#### Question 3: Alpha and Omega II
What was the month and day of the last sale? Store as a tuple in that order and assign the tuple to the variable `last_date`.

<!-- grade_id: cell-a6304e08d438b5d0, points: 0, grade: false, solution: true -->
```python
# your code here
mth = df_cleaned.loc[38222]["trans_month"]
day = df_cleaned.loc[38222]["trans_day"]
last_date = (mth,day)
last_date
```

Output:
```
(6, 30)
```

<!-- grade_id: cell-5ef9e6fb442b7a99, points: 4, grade: true, solution: false -->
```python
# Q3 Test Cases
assert len(last_date) == 2
assert str(last_date[0]).isnumeric() 
assert str(last_date[1]).isnumeric() 
```

#### Question 4: Cats vs Dogs

Which animal product type is most popular?

<!-- grade_id: cell-10183753169be5e8, points: 0, grade: false, solution: true -->
```python
# your code here
df_cleaned.prod_animal_type.value_counts()
names = df_cleaned.prod_animal_type.value_counts().keys()
most_pop = names[0]
```

<!-- grade_id: cell-996a95ebb5eec26f, points: 4, grade: true, solution: false -->
```python
# Q4 Test Cases
assert type(most_pop) == str
assert most_pop in ['cat', 'dog']
```

#### Question 5: More Money More Problems I

What was the total dollar amount made in the month of January? Store this in the variable `jan_rev`.

<!-- grade_id: cell-c51d5edcf4fe931c, points: 0, grade: false, solution: true -->
```python
# your code here
sales_jan = df_cleaned[df_cleaned["trans_month"] == 1]
jan_rev = sales_jan["total_sales"].sum()
```

<!-- grade_id: cell-9a89e109181bc6b0, points: 4, grade: true, solution: false -->
```python
# Q5 Test Cases
assert  40000 <= jan_rev <= 60000
```

#### Question 6: More Money More Problems II

What was the total dollar amount made in the month of January? Store this in the variable `june_rev`.

<!-- grade_id: cell-0f1521ee38985bdd, points: 0, grade: false, solution: true -->
```python
# your code here
sales_june = df_cleaned[df_cleaned["trans_month"] == 6]
june_rev = sales_june["total_sales"].sum()
```

<!-- grade_id: cell-5ac7a212f71fe758, points: 4, grade: true, solution: false -->
```python
# Q6 Test Cases
assert  500000 <= june_rev <= 600000
```

#### Question 7: Transaction Size

What is the average number of items bought in each transaction? Sore this in the variable `avg_num_items`.

<!-- grade_id: cell-aeccee5f203a3f64, points: 0, grade: false, solution: true -->
```python
# your code here

avg_num_items = df_cleaned['trans_quantity'].sum()/df_cleaned['trans_id'].drop_duplicates().count()
```

<!-- grade_id: cell-9107bb862558578a, points: 4, grade: true, solution: false -->
```python
# Q7 Test Cases
assert  0 <= avg_num_items <= 3
```

#### Question 8: Best Products I

What are the top ten product titles by the number of sales? Display in descending order. Store in variable `top_num_sales`.

<!-- grade_id: cell-6ebd3fdb04804030, points: 0, grade: false, solution: true -->
```python
# your code here

top_num_sales = df_cleaned['prod_title'].value_counts().head(10)
```

<!-- grade_id: cell-5cd99f7a3013a125, points: 4, grade: true, solution: false -->
```python
# Q8 Test Cases
assert len(top_num_sales) == 10
```

#### Question 9: Best Products II

What are the top ten product titles by total dollar amount made? Display in descending order. Store in variable `top_tot_sales`.

<!-- grade_id: cell-18c6a06ec971a5d9, points: 0, grade: false, solution: true -->
```python
# your code here

top_num_sales = df_cleaned.groupby('prod_title').agg('total_sales').sum().reset_index()
top_tot_sales =top_num_sales.sort_values("total_sales",ascending=False).head(10)
top_tot_sales=top_tot_sales["prod_title"]
top_tot_sales
```

Output:
```
11           Reddy Beddy
2               Cat Cave
8          Kitty Climber
15       Snoozer Hammock
14    Snoozer Essentails
20         Yum Fish-Dish
12         Scratchy Post
5         Feline Fix Mix
7            Foozy Mouse
18          Tuna Tasties
Name: prod_title, dtype: object
```

<!-- grade_id: cell-9a381e47ff024ad4, points: 4, grade: true, solution: false -->
```python
# Q9 Test Cases
assert len(top_tot_sales) == 10
```

#### Question 10: Bonus

What is the proportion of returning customers? Store as variable `prop_returning`.

<!-- grade_id: cell-59c5307f420fe231, points: 0, grade: false, solution: true -->
```python
# your code here
df2 = pd.DataFrame(df_cleaned.groupby("cust_id")["trans_id"].count())
returncustomers = (df2["trans_id"] > 1).sum()
prop_returning = returncustomers / len(df2)
```

<!-- grade_id: cell-84d62b95c61208f4, points: 0, grade: true, solution: false -->
```python
# Q10 Test Cases
assert 0 <= prop_returning <= 1
```