# Part III: Explore

![alt text](data/inu_neko_logo_small.png "Inu + Neko")

Hello! We are Inu + Neku and we are a Dog & Cat services and supplies store located in New York City. We just started our e-commerce business and need your help analyzing our data!

## Description

We need to make sure the data is clean before starting your analysis. As a reminder, we should check for:

- Duplicate records
- Consistent formatting
- Missing values
- Obviously wrong values

```python
import pandas as pd
import numpy as np
import matplotlib.pyplot as plt
```

```python
df_cleaned = pd.read_csv('data/inu_neko_orderline_clean.csv')
df_cleaned
```

```python
df_cleaned.info()
```

#### Question 1: Number of Orders

How many transactions are there?

<!-- grade_id: cell-976c898e7f2b44c4, points: 0, grade: false, solution: true -->
```python
# your code here

num_trans
```

<!-- grade_id: cell-25ca64f53b142f34, points: 4, grade: true, solution: false -->
```python
# Q1 Test Cases
assert 39000 > num_trans > 25000
```

#### Question 2: Alpha and Omega I
What was the month and day of the first sale? Store as a tuple in that order and assign the tuple to the variable `first_date`.

<!-- grade_id: cell-bd2b89a8035a9ab8, points: 0, grade: false, solution: true -->
```python
# your code here

first_date
```

<!-- grade_id: cell-8a2534037c10eb71, points: 4, grade: true, solution: false -->
```python
# Q2 Test Cases
assert len(first_date) == 2
assert str(first_date[0]).isnumeric() 
assert str(first_date[1]).isnumeric() 
```

#### Question 3: Alpha and Omega II
What was the month and day of the last sale? Store as a tuple in that order and assign the tuple to the variable `last_date`.

<!-- grade_id: cell-a6304e08d438b5d0, points: 0, grade: false, solution: true -->
```python
# your code here

last_date
```

<!-- grade_id: cell-5ef9e6fb442b7a99, points: 4, grade: true, solution: false -->
```python
# Q3 Test Cases
assert len(last_date) == 2
assert str(last_date[0]).isnumeric() 
assert str(last_date[1]).isnumeric() 
```

#### Question 4: Cats vs Dogs

Which animal product type is most popular?

<!-- grade_id: cell-10183753169be5e8, points: 0, grade: false, solution: true -->
```python
# your code here

most_pop
```

<!-- grade_id: cell-996a95ebb5eec26f, points: 4, grade: true, solution: false -->
```python
# Q4 Test Cases
assert type(most_pop) == str
assert most_pop in ['cat', 'dog']
```

#### Question 5: More Money More Problems I

What was the total dollar amount made in the month of January? Store this in the variable `jan_rev`.

<!-- grade_id: cell-c51d5edcf4fe931c, points: 0, grade: false, solution: true -->
```python
# your code here

jan_rev
```

<!-- grade_id: cell-9a89e109181bc6b0, points: 4, grade: true, solution: false -->
```python
# Q5 Test Cases
assert  40000 <= jan_rev <= 60000
```

#### Question 6: More Money More Problems II

What was the total dollar amount made in the month of January? Store this in the variable `june_rev`.

<!-- grade_id: cell-0f1521ee38985bdd, points: 0, grade: false, solution: true -->
```python
# your code here

june_rev
```

<!-- grade_id: cell-5ac7a212f71fe758, points: 4, grade: true, solution: false -->
```python
# Q6 Test Cases
assert  500000 <= june_rev <= 600000
```

#### Question 7: Transaction Size

What is the average number of items bought in each transaction? Sore this in the variable `avg_num_items`.

<!-- grade_id: cell-aeccee5f203a3f64, points: 0, grade: false, solution: true -->
```python
# your code here

avg_num_items
```

<!-- grade_id: cell-9107bb862558578a, points: 4, grade: true, solution: false -->
```python
# Q7 Test Cases
assert  0 <= avg_num_items <= 3
```

#### Question 8: Best Products I

What are the top ten product titles by the number of sales? Display in descending order. Store in variable `top_num_sales`.

<!-- grade_id: cell-6ebd3fdb04804030, points: 0, grade: false, solution: true -->
```python
# your code here

top_num_sales
```

<!-- grade_id: cell-5cd99f7a3013a125, points: 4, grade: true, solution: false -->
```python
# Q8 Test Cases
assert len(top_num_sales) == 10
```

#### Question 9: Best Products II

What are the top ten product titles by total dollar amount made? Display in descending order. Store in variable `top_tot_sales`.

<!-- grade_id: cell-18c6a06ec971a5d9, points: 0, grade: false, solution: true -->
```python
# your code here

top_tot_sales
```

<!-- grade_id: cell-9a381e47ff024ad4, points: 4, grade: true, solution: false -->
```python
# Q9 Test Cases
assert len(top_tot_sales) == 10
```

#### Question 10: Bonus

What is the proportion of returning customers? Store as variable `prop_returning`.

<!-- grade_id: cell-59c5307f420fe231, points: 0, grade: false, solution: true -->
```python
# your code here

prop_returning
```

<!-- grade_id: cell-84d62b95c61208f4, points: 0, grade: true, solution: false -->
```python
# Q10 Test Cases
assert 0 <= prop_returning <= 1
```