# Lab: Functions

In this lab we will go over:

- Help functions
- String functions
- Math functions
- Writing functions

#### Question 1: Help Me, Jupyter, You're My Only Hope

Try using the `help()` function. An example is provided below. Replace `sum` with a couple of other built in functions.

```python
help(sum)
```

Output:
```
Help on built-in function sum in module builtins:

sum(iterable, start=0, /)
    Return the sum of a 'start' value (default: 0) plus an iterable of numbers
    
    When the iterable is empty, return the start value.
    This function is intended specifically for use with numeric values and may
    reject non-numeric types.
```

#### Question 2: Tokenizer

Use the `Split()` function to seperate the following text into a `list` of individual words. Assign the output to the variable `words`.

<!-- grade_id: cell-969415462442da83, points: 0, grade: false, solution: true -->
```python
text = "We’ll always have Paris. We didn’t have, we, we lost it until you came to Casablanca."
# your code here

words =text.split()
```

<!-- grade_id: cell-f45c8cc8c15d26fb, points: 4, grade: true, solution: false -->
```python
# Q2 Test Cases
assert len(words) == 16
```

#### Question 3: Upper East Side

Use the `upper` function to convert the following sentence into all uppercase characters. Assign the output to the variable `upper_sent`.

<!-- grade_id: cell-81fac5d3c12a1219, points: 0, grade: false, solution: true -->
```python
sentence = 'a martini shaken not stirred'
# your code here

upper_sent = sentence.upper()
```

<!-- grade_id: cell-bf5690ad798770c2, points: 4, grade: true, solution: false -->
```python
# Q3 Test Cases
assert len(upper_sent) == 28
```

#### Question 4: Smallest Number I

Use the `Min()` function to find the smallest number from the following `list` of numbers. Assign the output to the variable `min_list`.

<!-- grade_id: cell-abc06e86240c9514, points: 0, grade: false, solution: true -->
```python
num_lst = [91.94, 40.37, 70.55, -89.28, 58.77,
           -37.47, -74.64, 85.49, 6.21, 5.55]
# your code here

min_list = min(num_lst)
```

<!-- grade_id: cell-9f349f38e230be86, points: 4, grade: true, solution: false -->
```python
# Q4 Test Cases
assert type(min_list) == float
```

#### Question 5: Smallest Number II

Write the function `minimum` that finds the smallest number of a list. DO NOT use any built in functions.

<!-- grade_id: cell-25cc10e14cf31404, points: 0, grade: false, solution: true -->
```python
def minimum(lst):
    # your code here
    lst.sort()
    return lst[0]
```

<!-- grade_id: cell-4991b01021403e5b, points: 4, grade: true, solution: false -->
```python
# Q5 Test Cases
assert minimum([0, 1, 2]) == 0
```

#### Question 6: Averages Assemble

Write the function `average` that finds the average of a list using a for loop. DO NOT use any built in functions. Remember the average is the sum of all the numbers in a list divided by the total number of items in the list.

<!-- grade_id: cell-b36eb5b8515677cc, points: 0, grade: false, solution: true -->
```python
def average(lst):
    # your code here
    sum = 0
    for i in lst:
        sum += i
        average = sum/len(lst)
    return average
```

<!-- grade_id: cell-d4f3bdc3cb80038f, points: 4, grade: true, solution: false -->
```python
# Q6 Test Cases
assert average([-1, 0, 1]) == 0
```

#### Question 7: Even

Write a function `even_numbers` that uses a for loop and outputs a list that contains all the even numbers from $n_1$ inclusive up to $n_2$ exclusive.

<!-- grade_id: cell-0bc5bde282a3d8e8, points: 0, grade: false, solution: true -->
```python
def even_numbers(n1, n2):
    
    """Returns a list of all the even numbers from n1 inclusive up to n2 exclusive.

  Args:
    n1: The start number (inclusive).
    n2: The end number (exclusive).

  Returns:
    A list of all the even numbers from n1 inclusive up to n2 exclusive.
    """

    even_numbers = []
    for i in range(n1, n2):
        if i % 2 == 0:
            even_numbers.append(i)
    return even_numbers
```

```python
even_numbers(-2, 4)
```

Output:
```
[-2, 0, 2]
```

<!-- grade_id: cell-44f1ceff001d4534, points: 4, grade: true, solution: false -->
```python
# Q6 Test Cases
assert even_numbers(-2, 4) == [-2, 0, 2]
```

```python

```