# Lab: Intro to Jupyter and Python

In this lab we will go over:

- Navigating jupyter notebooks
- Types (e.g. str, int, float, bool, NoneType)
- Operators
- Variables
- Data Structures

#### Question 1: And So It Begins

Bellow is your first code cell. Try running it by clicking on the cell and then the `Run` button above.

Alternatively you can also use the hotkeys `shift` + `enter`. 


```python
"Hello world."
```

Output:
```
'Hello world.'
```

Try running this next cell.

```python
# This is a comment. 
# Python will never evaluate comments.
1
2
```

Output:
```
2
```

Notice how only the 2 is in the output. If you wanted to display multiple things from a cell, try using the `print()` function.

```python
# Everything prints.
print(1)
print(2)
```

Output:
```
1
2
```

#### Question 2: Basic Types

Here's a refresher on some of the basic types in python. We've provided two examples of each. Run the cells below.

```python
"Hello"
```

Output:
```
'Hello'
```

```python
# You can see the type of a value by using type()
type('world')
```

Output:
```
str
```

```python
0
```

Output:
```
0
```

```python
type(1)
```

Output:
```
int
```

```python
1.0
```

Output:
```
1.0
```

```python
type(3.1415)
```

Output:
```
float
```

```python
True
```

Output:
```
True
```

```python
type(False)
```

Output:
```
bool
```

```python
None
```

```python
type(None)
```

Output:
```
NoneType
```

#### Question 3: Casting Call

We can convert types from one to another. This is called casting. Run the cells below.

```python
str(1)
```

Output:
```
'1'
```

```python
str(0.0)
```

Output:
```
'0.0'
```

```python
int('1')
```

Output:
```
1
```

```python
int('2')
```

Output:
```
2
```

Sometimes it doesn't work. For instance, you can't convert characters from the alphabet into an `int`.

```python
# Error
int('A')
```

Output:
```
ValueError                                Traceback (most recent call last)
<ipython-input-18-1f3a50436274> in <module>
      1 # Error
----> 2 int('A')

ValueError: invalid literal for int() with base 10: 'A'
```

#### Question 4: Operators
In Python we can apply mathematical operators to numbers much like a calculator.
- `+` addition
- `-` subtraction
- `*` multiplication
- `**` exponentiation
- `/` division

We can also use parentheses to group operations together. Mathematical operators are evaluated in PEMDAS order (Parentheses, Exponentiation, Multiplication / Division, Addition / Subtraction).

Here are some examples:

```python
2 + 1
```

Output:
```
3
```

```python
1 / 3
```

Output:
```
0.3333333333333333
```

We can make it more complicated. Notice how we can combine floats and integers together.

```python
(1 + 1) / 3.14 ** 2
```

Output:
```
0.2028479857195018
```

We can also apply some mathematical operators to `str`.

```python
# Concatentation
'hello' + ' ' + 'world'
```

Output:
```
'hello world'
```

```python
# Repetition
'hello' * 3
```

Output:
```
'hellohellohello'
```

But not all of them work.

```python
# Error
'hello' * 'hello'
```

Output:
```
TypeError                                 Traceback (most recent call last)
<ipython-input-24-267e01d06f65> in <module>
      1 # Error
----> 2 'hello' * 'hello'

TypeError: can't multiply sequence by non-int of type 'str'
```

```python
# Error
'hello' / 2
```

Output:
```
TypeError                                 Traceback (most recent call last)
<ipython-input-25-cc9b36dc3278> in <module>
      1 # Error
----> 2 'hello' / 2

TypeError: unsupported operand type(s) for /: 'str' and 'int'
```

#### Question 5: Creating a New Code Cell

Try writing your own arithmetic expressions. First create a new code cell. Click on the cell below and navigate to `Insert` -> `Insert Cell Below`.

You can also press `esc` and then press `B`. If you are interested in more hotkeys press `esc` and then press `H`. Also, here is a more comprehensive guide on Jupyter Notebook [usage](https://nbviewer.jupyter.org/github/ipython/ipython/blob/3.x/examples/Notebook/Notebook%20Basics.ipynb).

```python
# Create some new code blocks
```

```python

```

#### Question 6: Variables I

Assign the variable `name` to a `str`.

<!-- grade_id: cell-138c46f3d39d8a2b, points: 0, grade: false, solution: true -->
```python
# your code here

name = "myname"
```

Verify that `name` is a `str` using `type()`. 

<!-- grade_id: cell-10c9edff8f514451, points: 0, grade: false, solution: true -->
```python
# your code here
type(name)
```

Output:
```
str
```

The cell bellow contains test cases. Some are visible and some are not. These are used to autograde each coding assignment. Use them as a sanity check. 
<!--
To get full credit all the test cases must pass.
-->

<!-- grade_id: cell-e45aa81bee87e649, points: 1, grade: true, solution: false -->
```python
# Q6 Test Cases
assert name is not None
```

#### Question 7: Variables II
Assign the variable `truth` to a `bool`.

<!-- grade_id: cell-0a93e70eeef58c7e, points: 0, grade: false, solution: true -->
```python
# your code here

truth = True
```

Verify that `truth` is a `bool` using `type()`

<!-- grade_id: cell-1f91e533135f0c74, points: 0, grade: false, solution: true -->
```python
# your code here
type(truth)
```

Output:
```
bool
```

<!-- grade_id: cell-c664f0e8c605e8f4, points: 1, grade: true, solution: false -->
```python
# Q7 Test Cases
assert truth is not None
```

#### Question 8: Lists I

Create a `list` which contains 4 different colors. Call it `color_list`.

<!-- grade_id: cell-0ea547a6e7ae282e, points: 0, grade: false, solution: true -->
```python
# your code here

color_list = ['blue','red','yellow','green']
```

<!-- grade_id: cell-6e72a614ff936068, points: 1, grade: true, solution: false -->
```python
# Q8 Test Cases
assert type(color_list) == list
```

#### Question 9: Lists II
Access the first color of your `color_list`. Store it as the variable `first`.

<!-- grade_id: cell-7732ff56a84bb907, points: 0, grade: false, solution: true -->
```python
# your code here

first = color_list[0]
```

<!-- grade_id: cell-8516b05f5173c907, points: 1, grade: true, solution: false -->
```python
# Q9 Test Cases
assert type(first) == str
```

#### Question 10: Lists III
Access the last color of your `color_list`. Store it as the variable `last`.

<!-- grade_id: cell-72f8c5811ef60991, points: 0, grade: false, solution: true -->
```python
# your code here

last=color_list[3]
```

<!-- grade_id: cell-5b270e4755bb18a8, points: 1, grade: true, solution: false -->
```python
# Q10 Test Cases
assert type(last) == str
```

#### Question 11: Lists IV
Access the first 2 colors of your `color_list`. Store it as the variable `first_two`.

<!-- grade_id: cell-fdc024b7caecc40a, points: 0, grade: false, solution: true -->
```python
# your code here

first_two = color_list[0:2]
```

<!-- grade_id: cell-0cd6abb1240378a4, points: 1, grade: true, solution: false -->
```python
# Q11 Test Cases
assert type(first_two) == list
assert len(first_two) == 2
```

#### Question 12: Lists V
Access the last 2 colors of your `color_list`. Store it as the variable `last_two`.

<!-- grade_id: cell-f278c17eb74a1732, points: 0, grade: false, solution: true -->
```python
# your code here

last_two = color_list[:-2]
```

<!-- grade_id: cell-a891b0369a1612cd, points: 1, grade: true, solution: false -->
```python
# Q12 Test Cases
assert type(last_two) == list
assert len(last_two) == 2
```

#### Question 13: Lists VI
Change the second item of your to the color 'cyan'.

<!-- grade_id: cell-49aaa791fc5d51b5, points: 0, grade: false, solution: true -->
```python
# your code here

color_list[1] = 'cyan'

color_list
```

Output:
```
['blue', 'cyan', 'yellow', 'green']
```

<!-- grade_id: cell-f3614744785f71c0, points: 1, grade: true, solution: false -->
```python
# Q13 Test Cases
assert type(color_list) == list
```

#### Question 14: Lists VII
Here is another `list`. Notice how there can be different types in the `list`! Access the first to the third item in the `random_list`. Store it as the variable `first_to_third`.

<!-- grade_id: cell-19f842f1d20fbaab, points: 0, grade: false, solution: true -->
```python
random_list = ['Five', 9, False, 4, 3.14, -10, -1, 8, 4.0, -7]
# your code here

first_to_third = random_list [0:3]
```

<!-- grade_id: cell-8be634109cd3518f, points: 1, grade: true, solution: false -->
```python
# Q14 Test Cases
assert type(first_to_third) == list
assert len(first_to_third) == 3
```

#### Question 15: Lists VIII
Access the fifth item to the second to last item in `random_list`. Store it as the variable `second_to_last`.

<!-- grade_id: cell-0edfe19497391bf4, points: 0, grade: false, solution: true -->
```python
# your code here

second_to_last = random_list[5:]
```

<!-- grade_id: cell-118580e9244e2ae9, points: 1, grade: true, solution: false -->
```python
# Q15 Test Cases
assert type(second_to_last) == list
assert len(second_to_last) == 5
```

#### Question 16: Lists IX
We can show how many elements are in a `list` by using `len()`. Store the length of `random_list` as the variable `rand_len`.

<!-- grade_id: cell-a72c1a51bd6cfe68, points: 0, grade: false, solution: true -->
```python
# your code here

rand_len = len(random_list)
```

<!-- grade_id: cell-50434ef3f05d9494, points: 1, grade: true, solution: false -->
```python
# Q16 Test Cases
assert type(rand_len) == int
```

#### Question 17: Lists X
Add a new number to the end of `random_list` using the `.append()` function.

<!-- grade_id: cell-ed22cebf97df8678, points: 0, grade: false, solution: true -->
```python
# your code here
random_list.append(0)
```

<!-- grade_id: cell-7a61103cf542d038, points: 1, grade: true, solution: false -->
```python
# Q17 Test Cases
assert type(random_list) == list
```

#### Question 18: Dictionaries I
Consider the table below:

| Name | Favorite Food |
| ---- | ------------ |
| Williams | Pizza |
| Osaka | Salad |
| Raducanu | Curry |
| Fernandez | Poutine |

Create a `dict` that takes in *Name* as the key and *Favorite Food* as the value. Store it as the variable `food_dict`.

<!-- grade_id: cell-6d408ba0a95d5c48, points: 0, grade: false, solution: true -->
```python
# your code here

food_dict = {"Williams": "Pizza", "Osaka" : "Salad", "Raducanu" : "Curry", "Fernandez": "Poutine"}
```

<!-- grade_id: cell-e17c01b211264ebc, points: 1, grade: true, solution: false -->
```python
# Q18 Test Cases
assert type(food_dict) == dict
```

#### Question 19: Dictionaries II
Add your name and your favorite food to the same `dict`.

<!-- grade_id: cell-69c71b8bd8a9b056, points: 0, grade: false, solution: true -->
```python
# your code here
food_dict["dennis"] = "fish"
```

```python
food_dict
```

Output:
```
{'Williams': 'Pizza',
 'Osaka': 'Salad',
 'Raducanu': 'Curry',
 'Fernandez': 'Poutine',
 'dennis': 'fish'}
```

<!-- grade_id: cell-dadd23c2544ff91a, points: 1, grade: true, solution: false -->
```python
# Q19 Test Cases
assert type(food_dict) == dict
```

#### Question 20: Dictionaries III
Use the `.values()` method to get the values from `food_dict`. Store it as the variable `food_values`.

<!-- grade_id: cell-f419c3109e6d54cc, points: 0, grade: false, solution: true -->
```python
# your code here

food_values = food_dict.values()
```

<!-- grade_id: cell-7df0d661d21de903, points: 1, grade: true, solution: false -->
```python
# Q20 Test Cases

```

#### Question 21: Dictionaries IV
Consider the table below:

| int_num | str_num |
| ---- | ------------ |
| 0| zero |
| 1 | one |
| 2 | two |
| 3 | three |

Create a `dict` that takes in *integer* (int) as the key and *string* (str) as the value. Store it as the variable `num_mapper`. Notice that the keys are integer numbers, not strings, (e.g. `0`, `1`, ... and not `"0"`, `"1"`, ... ) and the values are strings (`'zero'`, `'one'`, ... ).

<!-- grade_id: cell-4b56805caa6269c5, points: 0, grade: false, solution: true -->
```python
# your code here

num_mapper = {0: "zero", 1: "one", 2:"two", 3: "three"}
```

<!-- grade_id: cell-ecf0726e406a8860, points: 1, grade: true, solution: false -->
```python
# Q21 Test Cases

assert type(num_mapper) == dict
```

#### Question 22: Dictionaries IV
Access the value associated with the key `0`. Store it as the variable `zero`.

<!-- grade_id: cell-5a083ceff9281d7b, points: 0, grade: false, solution: true -->
```python
# your code here

zero = num_mapper[0]
```

<!-- grade_id: cell-cb1c9f351ccad292, points: 1, grade: true, solution: false -->
```python
# Q22 Test Cases

assert type(zero) == str
```

```python

```