# Part II: Scrubbing

![alt text](data/inu_neko_logo_small.png "Inu + Neko")

Hello! We are Inu + Neku and we are a Dog & Cat services and supplies store located in New York City. We just started our e-commerce business and need your help analyzing our data!

## Description

We need to make sure the data is clean before starting your analysis. As a reminder, we should check for:

- Duplicate records
- Consistent formatting
- Missing values
- Obviously wrong values

```python
import pandas as pd
import numpy as np
import matplotlib.pyplot as plt
```

```python
df = pd.read_csv('data/inu_neko_orderline.csv')
df
```

```python
df.info()
```

#### Question 1: Duplicate Records

How many duplicate transaction records are there? Do not count the first occurence as a duplicate. Assign your answer to the variable `dup_rows`.

<!-- grade_id: cell-098e0ee8e6444df1, points: 0, grade: false, solution: true -->
```python
# your code here

dup_rows
```

<!-- grade_id: cell-1b3c442ab9478645, points: 2, grade: true, solution: false -->
```python
# Q1 Test Cases
assert 500 >= dup_rows >= 300
```

#### Question 2: Drop Duplicate Records

Drop the duplicated records and assign the output `DataFrame` to a new variable called `df_cleaned`. Use this from now on.

<!-- grade_id: cell-3672260cd0ca48a1, points: 0, grade: false, solution: true -->
```python
# your code here

df_cleaned.head()
```

<!-- grade_id: cell-918643d81e7a8519, points: 2, grade: true, solution: false -->
```python
# Q2 Test Cases
assert 39000 >= len(df_cleaned) >= 38000
```

#### Question 3: Consistent Formatting

Fix the inconsistencies in the column `cust_state`.

<!-- grade_id: cell-294723711137757b, points: 0, grade: false, solution: true -->
```python
# your code here

df_cleaned.head()
```

<!-- grade_id: cell-cde65480600ea092, points: 4, grade: true, solution: false -->
```python
# Q3 Test Cases
assert 55 > df_cleaned['cust_state'].nunique() > 45
```

#### Question 4: Fistful of Dollars I

Validate the prices in `prod_price`. Remove any rows that are obviously wrong.

<!-- grade_id: cell-7aca7aae23c61ceb, points: 0, grade: false, solution: true -->
```python
# your code here

df_cleaned.head()
```

<!-- grade_id: cell-2fb81eaac500d459, points: 4, grade: true, solution: false -->
```python
# Q6 Test Cases
assert df_cleaned['prod_price'].max() < 500
assert df_cleaned['prod_price'].min() > 0
```

#### Question 5: Fistful of Dollars II

Correct the values in the `total_sales` column. Use `prod_price` and `trans_quantity`.

<!-- grade_id: cell-c976d5d940db9862, points: 0, grade: false, solution: true -->
```python
# your code here

```

<!-- grade_id: cell-3e6f21431145cbc5, points: 4, grade: true, solution: false -->
```python
# Q5 Test Cases
assert df_cleaned['total_sales'].dtype == float
assert np.all(df_cleaned['total_sales'] > 0)
```

#### Question 6: Missing Values I

How many missing values are there?

<!-- grade_id: cell-c1b6d55be1c36350, points: 0, grade: false, solution: true -->
```python
# your code here

num_nans
```

<!-- grade_id: cell-4f7ee5da10bd6b84, points: 4, grade: true, solution: false -->
```python
# Q6 Test Cases
assert 40000 > num_nans > 20000
```

#### Question 7: Missing Values II

Drop any columns with missing values.

<!-- grade_id: cell-6cbc4895cc6ed0e2, points: 0, grade: false, solution: true -->
```python
# your code here

df_cleaned.head()
```

<!-- grade_id: cell-97e3efb04b5340cb, points: 4, grade: true, solution: false -->
```python
# Q7 Test Cases
```