# Activity: Control Flow with Data Structures

## Introduction

In this activity, you will practice using control flow with data structures.
Control flow includes the following, and all or some of these may be used in this activity:
- `if` statements
- `for` loops
- `break` statements


#### Question 1

Use the created list `numbers` (which holds numbers from 1 to 10) to loop through `numbers`.
In the loop, and add each number squared to the `squares` list.

<!-- grade_id: cell-451980dfd7340435, points: 0, grade: false, solution: true -->
```python
# Your code here

numbers = [1,2,3,4,5,6,7,8,9,10]

squares = []
for num in numbers:
    squares.append(num ** 2)
```

<!-- grade_id: cell-404dd8af35136e63, points: 1, grade: true, solution: false -->
```python
# Question 1 Grading Checks

assert len(squares) == 10, "Double check the length your squares list."
```

#### Question 2

You are given the dictionary `students`, containing the names and their respective scores.

Create a new **list** called `top_students` which will hold **tuples** containing the name and score of each student **if the student's score is _greater than 90_**.


`top_students` will look similar to this:
```python
[('Alice', 95), ..., ('Emily', 100)]
```

```python
students = {
    'Alice': 95,
    'Bob': 80,
    'Charlie': 92,
    'David': 70,
    'Emily': 100
}
```

<!-- grade_id: cell-3b85f507bb917a41, points: 0, grade: false, solution: true -->
```python
# Your code here

top_students = [(name, score) for name, score in students.items() if score > 90]
top_students
```

Output:
```
[('Alice', 95), ('Charlie', 92), ('Emily', 100)]
```

<!-- grade_id: cell-dc908e4b76869453, points: 1, grade: true, solution: false -->
```python
# Question 2 Grading Checks

assert type(top_students) == list, "Double check the type of your top_students variable."
```

#### Question 3

You are given a list `exotic_birds`, containing the name and lifespan of various exotic birds.

Iterate through `exotic_birds` to create a dictionary called `higher_lifespans` that holds 
the name and lifespan of birds with a lifespan greater than 30 years.

`higher_lifespan`'s **keys** will the bird's name and the **values** will the bird's lifespan.

```python
exotic_birds = [
    ("Hyacinth Macaw", 70),
    ("African Grey Parrot", 50),
    ("Scarlet Macaw", 45),
    ("Cockatiel", 20),
    ("Kakapo", 75),
    ("Toucan", 20),
    ("Moluccan Cockatoo", 70),
    ("Lorikeet", 25),
    ("Crowned Pigeon", 25),
    ("Quetzal", 15)
]
```

<!-- grade_id: cell-5358005192964434, points: 0, grade: false, solution: true -->
```python
# Your code here

higher_lifespans = {name: lifespan for name, lifespan in exotic_birds if lifespan > 30}
```

<!-- grade_id: cell-175f68673bfd518c, points: 1, grade: true, solution: false -->
```python
# Question 3 Grading Checks

assert type(higher_lifespans) == dict, "Make sure higher_lifespans is a dictionary."

```

#### Question 4

You are given the list `numbers`.

By iterating through the list, create **two lists** `even_numbers` and `odd_numbers` to contain all the even & odd numbers of the list respectively.

```python
numbers = [1, 2, 3, 4, 5, 6, 7, 8, 9, 10]
```

<!-- grade_id: cell-a446bf04850bf677, points: 0, grade: false, solution: true -->
```python
# Your code here

even_numbers = []
odd_numbers = []

for num in numbers:
    if num % 2 == 0:
        even_numbers.append(num)
    else:
        odd_numbers.append(num)
```

<!-- grade_id: cell-9013dfbc0e8de0bb, points: 1, grade: true, solution: false -->
```python
# Question 4 Grading Checks

assert type(even_numbers) == list, "Double check the type of your even_numbers variable."
assert type(odd_numbers) == list, "Double check the type of your odd_numbers variable."
```

#### Question 5

You are given the list `sales` that contains the total purchases for each day of the week for a retail store.

Iterate through the list to create a variable `total_sales` that _accumulates_ the total sales of the week so far until the value of `total_sales` exceeds `1,000.00`. At that point, print ***"We hit our sales goal!"*** and stop iterating (`total_sales` will no longer be update).

```python
sales = [100.00, 200.00, 300.00, 400.00, 150.00, 250.00, 350.00, 150.00, 100.00, 250.00]

```

<!-- grade_id: cell-3ccba3572ec63f9b, points: 0, grade: false, solution: true -->
```python
# Your code here

total_sales = 0

for sale in sales:
    total_sales += sale
    if total_sales > 1000:
        print("We hit our sales goal!")
        break
```

Output:
```
We hit our sales goal!
```

<!-- grade_id: cell-e3e134dece0d5987, points: 1, grade: true, solution: false -->
```python
# Question 5 Grading Checks

assert type(total_sales) == float, "Double check the type of your total_sales variable."
assert total_sales < sum(sales), "Did you stop iterating when total_sales reached over 1,000.00?"

```

```python

```