# Activity: Example of a Typical Notebook on Coursera

## Introduction

In this activity you will practice using Jupyter Notebook on Coursera. This is
a simple activity to get you familiar with activities using Jupyter Notebook in
this course.

The next video for the course will be a recording walkthrough going over this
activity! So if you're feeling a bit stuck, feel free to look at the video
walkthrough first!

This activity will cover the following topics:
- Running a "frozen" cell
- Running an editable cell
- Using & running the "frozen" question checking cell 

#### Question 1

The cell below is "frozen" and therefore can't edited. However, you can (and
should) run the cell.

In the cell, you'll see Python "comments" marked with a `#` at the beginning
of the line. These lines are commonly used to help give context about the Python
code written. In these activities, these are usually direct and/or hints about
the question.

You'll also see some simple Python code. Again, you can't edit this cell. But
running this cell will still execute the code within the cell.

**For this question, simply _run the cell below_:**

```python
# This cell is "frozen" and can't be edited
# But this cell can be run

# Run this cell
greeting = 'Hello!'
print(greeting)
```

Output:
```
Hello!
```

After running the cell, there is a ***"Question Grading Checks"*** cell.
This cell comes after the work done for the question and is "frozen" (can't be
edited but can be run).

This cell has two purposes:
1. Running this cell, gives you a quick "check" to see if your code is "correct" 
   enough. 
   - If it is "correct", the cell should pretty much do "nothing".
   - However, if your code isn't quite right, it may trigger an error. You
     should change your code (possibly doing a kernal restart as well) and rerun
     your code and the ***"Question Grading Checks"*** cell.
2. There are sometimes "hidden checks" that will only be run when you submit the
   activity. These will run against your code.

**To check your question, simply _run the cell below_:**

<!-- grade_id: cell-699f774bb9a3c032, points: 1, grade: true, solution: false -->
```python
# Question 1 Grading Checks
# Note this cell is "frozen" and not editable, but can still be run

assert greeting == 'Hello!', 'Did you run the cell above with Python code?'
```

#### Question 2

The cell below is a normal cell and can be edited by you to put in your code to
answer the question.

This time, to show you what the ***"Question Grading Checks"*** cell looks like
when your code isn't "correct", you'll do the following:

1. **First** _run the cell below with **no edits**_
2. **Second** _run the next cell, the ***"Question Grading Checks"*** cell._
   - This will produce an error. This is what the cell's output looks like if
     your code is not "correct".
3. **Third** _edit the cell so it instead has `number_one = 1` and run the cell_
   - This is the correct answer.
4. **Lastly** _run the following cell, the ***"Question Grading Checks"*** cell._
   - If you ran the cell above the ***"Question Grading Checks"*** cell after
     editing the code, you should see ***no output*** from the ***"Question
     Grading Checks"*** cell.

In general, if the ***"Question Grading Checks"*** cell gives no error, your 
code is _likely_ correct.

<!-- grade_id: cell-38ed0b6042e4c7d4, points: 0, grade: false, solution: true -->
```python
# Your code here
# REMINDER:
## First: run this cell with no modifications
## Second: run the "Question Grading Checks" cell; this will give an error
## Third: edit and then run this cell so the code instead says: number_one = 1
## Lastly: run the "Question Grading Checks" cell; this will give no output
number_one = 1



```

<!-- grade_id: cell-17bac1e5eb98a8d3, points: 1, grade: true, solution: false -->
```python
# Question 2 Grading Checks

assert number_one == 1, 'Did you edit and run the cell so it has: number_one = 1'
```

#### Question 3

This final question will demonstrate what happens when you have code that 
_passes_ the ***"Question Grading Checks"*** cell when working on the activity
in the notebook, but _does not pass_ the "hidden checks" that are used to score
the assignment when it is submitted. But don't worry, you can always make
changes and resubmit to get the answer right!

To best demonstrate this, the cell below is editable but already has some code
in the cell. Specically, it has the code `this_question = 1`. This is actually
_incorrect_; the correct answer is actually `this_question = 3`.

Here are the steps you should do:
1. **First** _run the cell below with **no edits**_
2. **Second** _run the next cell, the ***"Question Grading Checks"*** cell._
   - This will _not_ produce an error and you should see ***no output*** from
     the ***"Question Grading Checks"*** cell. This is because the code passes
     the check from this cell.
   - However, there are _"hidden checks"_ that will only run when this
     activity assignment is submitted. If you followed along so far, this
     question will be marked _wrong_.
3. **Third** _submit the activity (at the top of the notebook page)_
   - This will submit the notebook and you can see the "grade" for your activity
     assignment. You should see that you will recieve 0 points for this
     question.
   - But don't worry! This won't affect you passing the course! If you want, you
     can reopen the assignment and correct the code to recieve full points.
4. **[OPTIONAL] Finally** _reopen the assignment, change the code for this
   question to `this_question = 3`. Then resubmit this activity assignment._
   - After submitting, you'll see a new grade for this assignment, and you
     should recieve the full points for this question!

In general, if the ***"Question Grading Checks"*** cell gives no error, your 
code is _likely_ correct. But there might be hidden checks.

***REMEMBER: You can always make changes and resubmit an activity assignment to get full
points***

<!-- grade_id: cell-b35f93ae2783dbb5, points: 0, grade: false, solution: true -->
```python
# Note this is the "wrong" answer, but submit anyway
this_question = 3
# After submitting at least once, you can get the correct answer by instead
# changing the code to look like this:
# this_question = 3



```

<!-- grade_id: cell-913d39e020f3a6be, points: 1, grade: true, solution: false -->
```python
# Question 3 Grading Checks

assert this_question > 0, 'Did you change the code for this_question = ... ?'
```

```python

```