# Activity: Using Conditionals

## Introduction

In this activity, you will use conditionals to solve the following questions.
This activity contains:
- `if` statements
- `if-else` statements
- `if-elif-else` statements
- Nested `if` statements
- `if` statements with logical operators


#### Question 1

Write an `if-else` statement that assigns the value `True` to the variable `is_even` if `number` is even or `False` if `number` is odd.

```python
number = 24
is_even = False
```

<!-- grade_id: cell-8d40892bb13719be, points: 0, grade: false, solution: true -->
```python
# Your code here

if number % 2 == 0:
    is_even = True
else:
    is_even = False
```

<!-- grade_id: cell-d99d403b48b1b9a9, points: 1, grade: true, solution: false -->
```python
# Question 1 Grading Checks

assert type(is_even) == bool
```

#### Question 2

Write an `if` statement that will assign the value 0 to the variable `is_negative` if `another_number` is negative.

```python
another_number = -27
is_negative = 1
```

<!-- grade_id: cell-98ec20804d9c03d7, points: 0, grade: false, solution: true -->
```python
# Your code here

if another_number < 0:
    is_negative = 0
```

<!-- grade_id: cell-76e229c364465aba, points: 1, grade: true, solution: false -->
```python
# Question 2 Grading Checks

assert type(another_number) == int
```

#### Question 3

Write an `if-elif-else` statement that assigns the string `"positive"` to the variable `sign` if the `integer_number` is positive, `"negative"` if `integer_number` is negative, or `"zero"` if `integer_number` is 0.

```python
integer_number = 0
sign = ''
```

<!-- grade_id: cell-4c1a2ee95c7312e4, points: 0, grade: false, solution: true -->
```python
# Your code here

if integer_number > 0:
    sign = 'positive'
elif integer_number < 0:
    sign = 'negative'
else:
    sign = 'zero'
```

<!-- grade_id: cell-3a2dc83f6248a011, points: 1, grade: true, solution: false -->
```python
# Question 3 Grading Checks

assert type(sign) == str
assert len(sign) > 0

```

#### Question 4

You are given a `dict` called `person` with a person's name and an age.

Write a nested `if` statement that first checks if the person is old enough to vote, and assigns the value `True` to the variable `can_vote` if they are older than 17.

Then, check if they are 21 or older and assigns the value `True` to the `can_rent_car` variable.

Finally, print the `can_vote` and `can_rent_car` variables.


```python
person = {
    "name": "James Dean",
    "age": 19,
}
can_vote = False
can_rent_car = False
```

<!-- grade_id: cell-233cad696fe6145b, points: 0, grade: false, solution: true -->
```python
# Your code here

if person["age"] > 17:
    can_vote = True
    if person["age"] >= 21:
        can_rent_car = True

print("Can vote:", can_vote)
print("Can rent car:", can_rent_car)
```

Output:
```
Can vote: True
Can rent car: False
```

<!-- grade_id: cell-552f00c9c8a389d6, points: 1, grade: true, solution: false -->
```python
# Question 4 Grading Checks

assert type(can_vote) == bool
assert type(can_rent_car) == bool

```

#### Question 5

Using the `and` operator and/or the `or` operator, write an `if` statement that assigns the value `True` to the variable `can_watch_movie` if the person is older than 13 and younger than 18, or if the person is older than 65.


```python
person = {
    "name": "Sheila Knight",
    "age": 86,
}

can_watch_movie = False
```

<!-- grade_id: cell-cc339b34e9b3cdb2, points: 0, grade: false, solution: true -->
```python
# Your code here

if (person["age"] > 13 and person["age"] < 18) or person["age"] > 65:
    can_watch_movie = True
```

<!-- grade_id: cell-da1e05d80203127f, points: 1, grade: true, solution: false -->
```python
# Question 5 Grading Checks

assert type(can_watch_movie) == bool
```

#### Question 6

Using the given dictionary `bird`, write a conditional statement of your choosing (`if`, `if-else`, `if-elif-else`, or nested `if` statements) that uses the properties of `bird` to assign the correct value to the `bird_type` variable, then print `bird_type`.

Here are possible bird types and their properties:
- cardinal: red, small, can fly, cannot swim
- ostrich: black, large, cannot fly, can swim
- purple martin: purple, small, can fly, can swim
- penguin: black, medium, cannot fly, can swim
- canary: yellow, small, can fly, cannot swim
- crow: black, small, can fly, cannot swim
- mallard: green, medium, can fly, can swim
- toucan: orange, medium, can fly, cannot swim
- unknown

```python
bird_colors = ["red", "blue", "yellow", "green", "purple", "orange", "black"]
bird_sizes = ["small", "medium", "large"]
bird_can_fly = [True, False]
bird_can_swim = [True, False]

bird = {
    'color': 'yellow',
    'size': 'small',
    'can_fly': True,
    'can_swim': False,
}

bird_type = ""
```

<!-- grade_id: cell-34ba91286fdac8af, points: 0, grade: false, solution: true -->
```python
# Your code here

if bird["color"] == "red" and bird["size"] == "small" and bird["can_fly"] and not bird["can_swim"]:
    bird_type = "cardinal"
elif bird["color"] == "black" and bird["size"] == "large" and not bird["can_fly"] and bird["can_swim"]:
    bird_type = "ostrich"
elif bird["color"] == "purple" and bird["size"] == "small" and bird["can_fly"] and bird["can_swim"]:
    bird_type = "purple martin"
elif bird["color"] == "black" and bird["size"] == "medium" and not bird["can_fly"] and bird["can_swim"]:
    bird_type = "penguin"
elif bird["color"] == "yellow" and bird["size"] == "small" and bird["can_fly"] and not bird["can_swim"]:
    bird_type = "canary"
elif bird["color"] == "black" and bird["size"] == "small" and bird["can_fly"] and not bird["can_swim"]:
    bird_type = "crow"
elif bird["color"] == "green" and bird["size"] == "medium" and bird["can_fly"] and bird["can_swim"]:
    bird_type = "mallard"
elif bird["color"] == "orange" and bird["size"] == "medium" and bird["can_fly"] and not bird["can_swim"]:
    bird_type = "toucan"

print(bird_type)
```

Output:
```
canary
```

<!-- grade_id: cell-68f3bc05d8c4071f, points: 0, grade: true, solution: false -->
```python
# Question 6 Grading Checks

assert type(bird_type) == str
assert len(bird_type) > 0

```

```python

```