# Activity: Using Dictionaries

## Introduction

In this activity you will practice the following:
- Creating dictionaries
- Accessing values in a dictionary
- Adding key-value pairs to a dictionary
- Modifying values in a dictionary
- Dictionary methods


#### Question 1

Create a dictionary called `student` which contains the following:
- name: derrick
- age: 17
- math_grade: 90

> Note: Make sure the keys and values match (such as spelling)

<!-- grade_id: cell-f2663b732b65da27, points: 0, grade: false, solution: true -->
```python
# Your code here

student = {
    "name": "derrick",
    "age": 17,
    "math_grade": 90
}

```

<!-- grade_id: cell-1c32794963a228b2, points: 1, grade: true, solution: false -->
```python
# Question 1 Grading Checks

assert isinstance(student, dict), 'Have you created a dictionary called student?'
assert len(student) == 3, 'Have you added the correct number of key-value pairs to the dictionary?'

```

#### Question 2

Access the value of the `math_grade` key in the `student` dictionary you created in Question 1 and assign it to a variable called `grade`.


<!-- grade_id: cell-eb29541abd47c277, points: 0, grade: false, solution: true -->
```python
# Your code here

grade = student["math_grade"]

```

<!-- grade_id: cell-4a3c1c086515c6ee, points: 1, grade: true, solution: false -->
```python
# Question 2 Grading Checks

assert isinstance(grade, int), 'Have you accessed the value of the math_grade key in the student dictionary?'

```

#### Question 3

Add a new key-value pair to the `student` dictionary you created in Question 1. The new key should be `english_grade` and the value should be `85`.


<!-- grade_id: cell-c1b0b1f08800416d, points: 0, grade: false, solution: true -->
```python
# Your code here

student["english_grade"] = 85

```

<!-- grade_id: cell-8fb6c1ea22c4b357, points: 1, grade: true, solution: false -->
```python
# Question 3 Grading Checks

assert len(student) == 4, 'Have you added the correct number of key-value pairs to the dictionary?'

```

#### Question 4

Modify the value of the `english_grade` key in the `student` dictionary you created in Question 1 to be `90`.


<!-- grade_id: cell-e4ac806bdcad6c21, points: 0, grade: false, solution: true -->
```python
# Your code here

student["english_grade"] = 90

```

<!-- grade_id: cell-83e07315196266d5, points: 1, grade: true, solution: false -->
```python
# Question 4 Grading Checks

assert len(student) == 4, 'Have you added the correct number of key-value pairs to the dictionary?'

```

#### Question 5

Using the given dictionary `top_student`, store all the values in the dictionary in a variable called `top_student_grades`.


```python
top_student = {
    "math_grade": 95,
    "english_grade": 98,
    "history_grade": 90,
    "science_grade": 93,
    "art_grade": 92,
    "music_grade": 96
}
```

<!-- grade_id: cell-9385357ef3f923e9, points: 0, grade: false, solution: true -->
```python
# Your code here

top_student_grades = list(top_student.values())

```

<!-- grade_id: cell-6e019848e1555855, points: 1, grade: true, solution: false -->
```python
# Question 5 Grading Checks

assert len(top_student_grades) == len(top_student), 'Have you created top_student_grades from the dictionary?'

```

#### Question 6

Using the given dictionary `top_student` from Question 5, store all the keys in the dictionary in a variable called `top_student_subjects`.


<!-- grade_id: cell-3a646b4d1ce21982, points: 0, grade: false, solution: true -->
```python
# Your code here

top_student_subjects = list(top_student.keys())

```

<!-- grade_id: cell-fd30e26df6dd8c06, points: 1, grade: true, solution: false -->
```python
# Question 6 Grading Checks

assert len(top_student_subjects) == len(top_student), 'Have you created top_student_subjects from the dictionary?'

```

#### Question 7

Using the given dictionary `top_student` from Question 5, store all the key-value pairs in the dictionary in a variable called `top_student_items`.


<!-- grade_id: cell-0f701edfb3110476, points: 0, grade: false, solution: true -->
```python
# Your code here

top_student_items = list(top_student.items())

```

<!-- grade_id: cell-4ef1798a88c7a5e9, points: 1, grade: true, solution: false -->
```python
# Question 7 Grading Checks

assert len(top_student_items) == len(top_student), 'Have you created top_student_items from the dictionary?'

```

```python

```