# Activity: Using Booleans

## Introduction

In this activity, you will use booleans to complete the following questions.
This activity includes:
- Logical operators
- Boolean expressions
- Membership operators


#### Question 1

Negate the variable `can_drive` and assign the result to a new variable called `cannot_drive`.


```python
can_drive = True
```

<!-- grade_id: cell-f1efd588c28f85e7, points: 0, grade: false, solution: true -->
```python
# Your code here
can_drive = True
cannot_drive = not can_drive

```

<!-- grade_id: cell-3e2c4b66749d726c, points: 1, grade: true, solution: false -->
```python
# Question 1 Grading Checks

assert isinstance(cannot_drive, bool), "Make sure you are assigning a boolean value to cannot_drive"

```

#### Question 2

Check if the variable `age` is greater than or equal to 18 and less than or equal to 25. Assign the result to a variable called `is_age_valid`.


```python
age = 20
```

<!-- grade_id: cell-ec4d89d103a07fc4, points: 0, grade: false, solution: true -->
```python
# Your code here
is_age_valid = 18 <= age <= 25
```

<!-- grade_id: cell-4ab58dea976d33ea, points: 1, grade: true, solution: false -->
```python
# Question 2 Grading Checks

assert isinstance(is_age_valid, bool), "Make sure you are assigning a boolean value to is_age_valid"

```

#### Question 3

Check if the variable `name` is in the list `names`. Assign the result to a variable called `is_name_valid`.

Note this means that `is_name_valid` will be **`True`** if `name` is **in** the list. Otherwise, `is_name_valid` will be assigned the value **`False`**.


```python
name = "John"
names = ["John", "Jane", "Jack"]
```

<!-- grade_id: cell-293ce88b675d6d5e, points: 0, grade: false, solution: true -->
```python
# Your code here

is_name_valid = name in name
```

<!-- grade_id: cell-5fd6417848f3ba30, points: 1, grade: true, solution: false -->
```python
# Question 3 Grading Checks

assert isinstance(is_name_valid, bool), "Make sure you are assigning a boolean value to is_name_valid"

```

#### Question 4

Check if the variable `movie` is **NOT** in the list `movies` and assign the result to a variable called `not_in_movies`.

Note this means that `not_in_movies` will be **`False`** if `movie` **is in** the list `movies`.
If `movie` is **NOT** in the list, `not_in_movies` will be assigned the value **`True`**.


```python
movie = "Tarzan"
movies = ["The Matrix", "Tangled"]
```

<!-- grade_id: cell-2972bb6150b318f2, points: 0, grade: false, solution: true -->
```python
# Your code here
not_in_movies = movie not in movies

```

<!-- grade_id: cell-4810d9d6bf3c166a, points: 1, grade: true, solution: false -->
```python
# Question 4 Grading Checks

assert isinstance(not_in_movies, bool), "Make sure you are assigning a boolean value to not_in_movies."

```

#### Question 5

Write a boolean expression that checks if the variable `whole_num` is not equal to `float_num`. Assign the result to a variable called `is_not_equal`.


```python
whole_num = 5
float_num = 5.0
```

<!-- grade_id: cell-e7025779b7e6ca4e, points: 0, grade: false, solution: true -->
```python
# Your code here

is_not_equal = whole_num != float_num
```

<!-- grade_id: cell-f2acd3be329f872b, points: 1, grade: true, solution: false -->
```python
# Question 5 Grading Checks

assert isinstance(is_not_equal, bool), "Make sure you are assigning a boolean value to is_not_equal"

```

#### Question 6

Write a boolean expression that checks if the variable `num` is greater than or equal to 10 and less than or equal to 20. Assign the result to a variable called `is_valid`.


```python
num = 15
```

<!-- grade_id: cell-4f9074c3ddc9723b, points: 0, grade: false, solution: true -->
```python
# Your code here

is_valid = 10 <= num <= 20
```

<!-- grade_id: cell-f31c0477c4ca5897, points: 0, grade: true, solution: false -->
```python
#### Question 6 Grading Checks

assert isinstance(is_valid, bool), "Make sure you are assigning a boolean value to is_valid"
```

```python

```