# Activity: Using Lists & Tuples

## Introduction

In this activity you will practice using and creating lists and tuples.
This activity contains the following topics:
- Creating lists
- Creating tuples
- Indexing
- Slicing
- List methods
- Unpacking


#### Question 1

Create a list called `countries` with the names of 5 different countries of your choice.


<!-- grade_id: cell-0614fe54fc242fa2, points: 0, grade: false, solution: true -->
```python
# Your code here
countries = ["a","b","c","d","e"]

print(countries)
```

Output:
```
['a', 'b', 'c', 'd', 'e']
```

<!-- grade_id: cell-3a329dfd53eab709, points: 1, grade: true, solution: false -->
```python
# Question 1 Grading Checks

assert isinstance(countries, list), 'Your variable is not a list'

```

#### Question 2

Using the `countries` list, assign the first and last countries in the list to the variables `first_country` and `last_country` respectively.


<!-- grade_id: cell-2d238064c27e485b, points: 0, grade: false, solution: true -->
```python
# Your code here

first_country = countries[0]

last_country = countries[-1]

print(first_country)

print(last_country)
```

Output:
```
a
e
```

<!-- grade_id: cell-10e571d348bfb921, points: 1, grade: true, solution: false -->
```python
# Question 2 Grading Checks

assert len(countries) == 5, 'Your list does not have 5 elements'
assert isinstance(first_country, str), 'Your variable is not a string'
assert isinstance(last_country, str), 'Your variable is not a string'

```

#### Question 3

Using the `countries` list, add a country to the end of the list.
Then assign the last 3 countries to a variable called `last_three_countries`.


<!-- grade_id: cell-09f04a4add4b7637, points: 0, grade: false, solution: true -->
```python
# Your code here


countries.append("f")

print(countries)

last_three_countries = countries[-3:]

last_three_countries
```

Output:
```
['a', 'b', 'c', 'd', 'e', 'f']
['d', 'e', 'f']
```

<!-- grade_id: cell-2067ef95fa07586c, points: 1, grade: true, solution: false -->
```python
# Question 3 Grading Checks

assert len(countries) == 6, 'Your list does not have 6 elements (one more element on the end)'
assert isinstance(last_three_countries, list), 'Your variable is not a list'
assert len(last_three_countries) == 3, 'Your list does not have 3 elements'

```

#### Question 4

Using the `countries` list, assign the **even** indexed countries to a variable called `even_countries`. Remember the **first index is 0** (an even number).

This means if your list was something like:
```python
countries = [
    'Lilliput',
    'Gilead',
    'Oceania',
    'Wakanda',
    'Genovia',
    'Panem',
]
```
Then `even_countries` will be:
```python
[
    'Lilliput',
    'Oceania',
    'Genovia',
]
```


<!-- grade_id: cell-4a5264eeab7a9c33, points: 0, grade: false, solution: true -->
```python
# Your code here

even_countries = countries[::2]

print(even_countries)

```

Output:
```
['a', 'c', 'e']
```

<!-- grade_id: cell-e753d72d2a6b5400, points: 1, grade: true, solution: false -->
```python
# Question 4 Grading Checks

assert isinstance(even_countries, list), 'Your variable is not a list'

```

#### Question 5

Add another country to the `countries` list and assign the length of the list to a variable called `num_countries`.


<!-- grade_id: cell-8162b18aba873a5b, points: 0, grade: false, solution: true -->
```python
# Your code here

countries.append("g")

print(countries)

num_countries = len(countries)

print(num_countries)
```

Output:
```
['a', 'b', 'c', 'd', 'e', 'f', 'g']
7
```

<!-- grade_id: cell-5d6cf2f545b99477, points: 1, grade: true, solution: false -->
```python
# Question 5 Grading Checks

assert isinstance(num_countries, int), 'Your variable is not an integer'
assert num_countries > 6, 'Check the number of countries; you should have only added 2 more countries after the initial 5'

```

#### Question 6

You're given a tuple containing superheroes, unpack the tuple into variables `superman`, `batman`, and `black_panther`
 (in the expected order).

```python
# Given tuple

superheroes = ('Superman', 'Batman', 'Black Panther')
```

<!-- grade_id: cell-a6d6bddc63d43806, points: 0, grade: false, solution: true -->
```python
# Your code here

superman, batman, black_panther = superheroes

print(superman), print(batman), print(black_panther)
```

Output:
```
Superman
Batman
Black Panther
(None, None, None)
```

<!-- grade_id: cell-adfc0a04f765e223, points: 1, grade: true, solution: false -->
```python
# Question 6 Grading Checks

assert isinstance(superman, str), 'The superman variable is not a string'
assert isinstance(batman, str), 'The batman variable is not a string'
assert isinstance(black_panther, str), 'The black_panther variable is not a string'

```

#### Question 7

Create a list of **5 tuples** called `planet_info` where each tuple contains a planet name, the planet's color, and the planet's number of moons.


For example, your list might look something like this:
```python
planet_info = [
    ('Coruscant', 'orange', 4),
    ...
    ('Krypton', 'green', 3),
]
```

<!-- grade_id: cell-5275e17599ccffcf, points: 0, grade: false, solution: true -->
```python
# Your code here

planet_info = [('Mercury', 'gray', 0), ('Venus', 'yellow', 0), ('Earth', 'blue', 1), ('Mars', 'red', 2), ('Jupiter', 'orange', 79)]


```

<!-- grade_id: cell-ac994f5244998529, points: 1, grade: true, solution: false -->
```python
# Question 7 Grading Checks

assert isinstance(planet_info, list), 'Your variable is not a list'
assert isinstance(planet_info[0], tuple), 'Your first item is not a tuple'
assert len(planet_info[0]) == 3, 'Your first tuple does not have 3 elements'
assert len(planet_info) == 5, 'Your list does not have 5 elements'

```

```python

```