# Activity: Using Iterators

## Introduction

In this activity you will practice writing `for` loops and using some built-in functions to iterate over lists and dictionaries.
This activity contains:
- `for` loops
- `range()`
- `values()`
- `items()`
- `keys()`

#### Question 1

You're given a `list` of numbers from 1 through 10 and a variable called `total`. Write a `for` loop that sums all the numbers in the list and then prints the `total`.


```python
numbers = [1, 2, 3, 4, 5, 6, 7, 8, 9, 10]
total = 0
```

<!-- grade_id: cell-bed85f309d313050, points: 0, grade: false, solution: true -->
```python
# Your code here

for number in numbers:
    total += number

print(total)
```

Output:
```
55
```

<!-- grade_id: cell-597af53664c994d5, points: 1, grade: true, solution: false -->
```python
# Question 1 Grading Checks

assert total > 0
```

#### Question 2

Write a `for` loop which uses the `range()` function to sum up the total of the first 100 numbers (1 through 100) and prints the `total`.


```python
total = 0
```

<!-- grade_id: cell-f7a55be450d84ab6, points: 0, grade: false, solution: true -->
```python
# Your code here

for number in range(1, 101):
    total += number

print(total)
```

Output:
```
5050
```

<!-- grade_id: cell-c747bcc5dc6a9e29, points: 1, grade: true, solution: false -->
```python
# Question 2 Grading Checks

assert total > 0, "Double check your range"
assert total < 6000, "Double check your range; did you reset total to 0 if you ran more than once?"

```

#### Question 3

Write a `for` loop which uses the `range()` function to sum up the total of the (positive) **even** numbers _less than 100_ and prints the `total`.

```python
total = 0
```

<!-- grade_id: cell-aee0a937e57bad33, points: 0, grade: false, solution: true -->
```python
# Your code here

for number in range(0, 100, 2):
    total += number

print(total)
```

Output:
```
2450
```

<!-- grade_id: cell-76aa2fc7272212d2, points: 1, grade: true, solution: false -->
```python
# Question 3 Grading Checks

assert total > 0, "Double check your range"
assert total < 2700, "Double check your range; did you reset total to 0 if you ran more than once?"

```

#### Question 4

Using the given `dict` of `people`, write a `for` loop that prints the name of each person and their age, and also sums **all** their ages and assigns it to the variable `combined_ages`.

```python
people = {
    "John": 32,
    "Rob": 23,
    "Jane": 27,
    "Sally": 29,
    "Bob": 35,
}

combined_ages = 0
```

<!-- grade_id: cell-bd0faea309531482, points: 0, grade: false, solution: true -->
```python
# Your code here

for name, age in people.items():
    print(f"{name} is {age} years old.")
    combined_ages += age

print(f"The combined ages of all people is {combined_ages}.")
```

Output:
```
John is 32 years old.
Rob is 23 years old.
Jane is 27 years old.
Sally is 29 years old.
Bob is 35 years old.
The combined ages of all people is 146.
```

<!-- grade_id: cell-c9225f7f44ebedb1, points: 1, grade: true, solution: false -->
```python
# Question 4 Grading Checks

assert combined_ages > 0
```

#### Question 5

Using the given `dict` of `star_wars_movies`, write a `for` loop that prints the name of each movie and adds the name of each movie to the `movies` `list`.

```python
star_wars_movies = {
    1: "A New Hope",
    2: "The Empire Strikes Back",
    3: "Return of the Jedi",
    4: "The Phantom Menace",
    5: "Attack of the Clones",
    6: "Revenge of the Sith",
    7: "The Force Awakens",
    8: "The Last Jedi",
    9: "The Rise of Skywalker",
}
movies = []
```

<!-- grade_id: cell-6de077036ec83470, points: 0, grade: false, solution: true -->
```python
# Your code here

for episode, title in star_wars_movies.items():
    print(title)
    movies.append(title)
```

Output:
```
A New Hope
The Empire Strikes Back
Return of the Jedi
The Phantom Menace
Attack of the Clones
Revenge of the Sith
The Force Awakens
The Last Jedi
The Rise of Skywalker
```

<!-- grade_id: cell-dbfdf9efeefb01cb, points: 1, grade: true, solution: false -->
```python
# Question 5 Grading Checks

assert len(movies) == 9, 'Did you include all the movie titles in the list?'
assert isinstance(movies[-1], str), 'Did you add just the movie title to the list?'

```

#### Question 6

Using the given `dict` of `popular_coding_languages`, write a `for` loop that adds the name of the **3 highest rated** languages to the `top_languages` `list` and prints the `top_languages` `list`.


```python
popular_coding_languages = {
    "Python": 8.2,
    "JavaScript": 8.0,
    "Java": 5.2,
    "C#": 4.7,
    "PHP": 2.7,
    "C++": 2.0,
    "C": 1.7,
    "Shell": 1.4,
    "Go": 1.4,
    "TypeScript": 1.3,
}
top_languages = []
```

<!-- grade_id: cell-fbe3830531f32cd1, points: 0, grade: false, solution: true -->
```python
# Your code here
for language, rating in sorted(popular_coding_languages.items(), key=lambda x: x[1], reverse=True)[:3]:
    top_languages.append(language)
    print(language)

print(f"The top 3 programming languages are {', '.join(top_languages)}.")
```

Output:
```
Python
JavaScript
Java
The top 3 programming languages are Python, JavaScript, Java.
```

<!-- grade_id: cell-70eb8bb525c91ac6, points: 1, grade: true, solution: false -->
```python
# Question 6 Grading Checks

assert len(top_languages) == 3

```

```python

```