# Activity: Variables in Python

## Introduction

In this activity you will practice the following:
- Assigning specific values / types to variables
- Overwriting variables

#### Question 1

Create variable `name` and assign any value you want to it.

<!-- grade_id: cell-70f715e9946944db, points: 0, grade: false, solution: true -->
```python
# Your code here
name = 0




```

<!-- grade_id: cell-8580037199148fe8, points: 1, grade: true, solution: false -->
```python
# Question 1 Grading Checks

print('Make sure you created the variable: name')
print(name)


```

Output:
```
Make sure you created the variable: name
0
```

#### Question 2

Overwrite the variable `name` with a string value of your name.


<!-- grade_id: cell-be343115dcfb85c2, points: 0, grade: false, solution: true -->
```python
# Your code here

name = "dennis"



```

<!-- grade_id: cell-7f05e9e7cbc01cd8, points: 1, grade: true, solution: false -->
```python
# Question 2 Grading Checks

assert isinstance(name, str), 'Have you overwritten the `name` variable to a string value of your name?'

```

#### Question 3

Assign the value of `2.7` to the variable `seconds`.

<!-- grade_id: cell-2b1965cc8b59bb23, points: 0, grade: false, solution: true -->
```python
# Your code here


seconds = 2.7


```

<!-- grade_id: cell-2301487d87568583, points: 1, grade: true, solution: false -->
```python
# Question 3 Grading Checks

assert isinstance(seconds, float), 'Have you assigned a float value to the variable seconds?'

```

#### Question 4

Assign the number of days in full week (an integer) to the variable `days_in_a_week`.

<!-- grade_id: cell-cbaf6bd647e91a4a, points: 0, grade: false, solution: true -->
```python
# Your code here


days_in_a_week = 7


```

<!-- grade_id: cell-028698e16746a872, points: 1, grade: true, solution: false -->
```python
# Question 4 Grading Checks

assert isinstance(days_in_a_week, int), 'Have you assigned an integer number to the days_in_a_week variable?'

```

#### Question 5

Assign boolean value `True` to the variable `python_is_fun`.

<!-- grade_id: cell-7f478f83233bdc01, points: 0, grade: false, solution: true -->
```python
# Your code here


python_is_fun = True


```

<!-- grade_id: cell-4cd79574ca080736, points: 1, grade: true, solution: false -->
```python
# Question 5 Grading Checks

assert isinstance(python_is_fun, bool), 'Have you assigned a boolean value to the python_is_fun variable?'
```

#### Question 6

Assign `None` (the `NoneType`) to the variable `nothing`.

<!-- grade_id: cell-36ea8cad9c22783e, points: 0, grade: false, solution: true -->
```python
# Your code here


nothing = None


```

<!-- grade_id: cell-0a155951c12b02b8, points: 1, grade: true, solution: false -->
```python
# Question 6 Grading Checks

assert nothing == None, 'Have you assigned the value None (the NoneType) to the nothing variable?'
```

```python

```