# Activity: Writing Functions

## Introduction

In this activity, you will practice writing functions. You will be asked to write functions that:
- Take in parameters
- Return values
- Print values
- Use loops
- Use conditional statements
- Use lists
- Use dictionaries
- Use tuples

All of the above topics have been covered in the previous lessons. If you are having trouble with any of these topics, please review the previous lessons before attempting this activity.

#### Question 1

Write a function called `greeting()` that should take in **one string parameter** and **prints** 'Hello there, `NAME`' where `NAME` is the parameter given.


Example _printed_ output of `greeting('Victor')`:
```python
'Hello there, Victor'
```

<!-- grade_id: cell-5ceeb05c28c416f8, points: 0, grade: false, solution: true -->
```python
# Complete the implementation of this function
def greeting(name):
    print(f'Hello there, {name}')


```

<!-- grade_id: cell-aaf4f64480fa88e8, points: 1, grade: true, solution: false -->
```python
# Question 1 Grading Checks

assert greeting, "Make sure you name your function 'greeting'."
assert greeting('Victor') is None, "Make sure your function doesn't return anything."
```

Output:
```
Hello there, Victor
```

#### Question 2

Write a function called `square()` that takes in **one number parameter** and **returns** the squared value of that number.

Example _returned_ output of `square(2)`:
```python
4
```

<!-- grade_id: cell-100ee93aeb9451fd, points: 0, grade: false, solution: true -->
```python
# Complete the implementation of this function
def square(num):
    return num ** 2


```

<!-- grade_id: cell-344ea701437acbba, points: 1, grade: true, solution: false -->
```python
# Question 2 Grading Checks

assert square, "Make sure you name your function 'square'."
assert square(2) == 4, "Make sure your function returns the correct value."
```

#### Question 3

Complete the implementation of the function `print_multiple()` that takes in **two parameters**, `string` and `num`, and **prints** `string` `num` times.

Example _printed_ output of `print_multiple('hello', 3)`:

```python
'hello'
'hello'
'hello'
```




<!-- grade_id: cell-4efcbaf7f6031e0b, points: 0, grade: false, solution: true -->
```python
# Complete the implementation of this function
def print_multiple(string, num):
    for i in range(num):
        print(string)


```

<!-- grade_id: cell-c2653be4db4e5954, points: 1, grade: true, solution: false -->
```python
# Question 3 Grading Checks

assert print_multiple, "Make sure you name your function 'print_multiple'."
assert print_multiple('hello', 3) is None, "Make sure your function doesn't return anything."
```

Output:
```
hello
hello
hello
```

#### Question 4

Write a function called `area_perimeter` that takes in **two number parameters**, `length` and `width`, and **returns a tuple** containing both the _area_ and _perimeter_ of a rectangle with the given `length` and `width`.

Example _returned_ output of `area_perimeter(2, 3)`:
```python
(6, 10)
```

<!-- grade_id: cell-2216f0a1d7febac1, points: 0, grade: false, solution: true -->
```python
# Complete the implementation of this function
def area_perimeter(length, width):
    area = length * width
    perimeter = 2 * (length + width)
    return (area, perimeter)

```

<!-- grade_id: cell-01e133b5fc8bd5e5, points: 1, grade: true, solution: false -->
```python
# Question 4 Grading Checks

assert area_perimeter, "Make sure you name your function 'area_perimeter'."
assert area_perimeter(2, 3) == (6, 10), "Make sure your function returns the correct tuple."

```

#### Question 5

Write a function called `dict_to_tuple_list()` that takes in **one dictionary parameter** where keys are the names of bond villans and the values are the movie they appeared. (See the defined `bond_villans` in the next cell).

The function then **returns** a _list of tuples_ where each `tuple` contains the name of the villain and the movie they appeared in.

Example of function:
```python
test_dictionary = {
    'Dr. No': 'Dr. No',
    'Mr. Big': 'Live and Let Die',
}

dict_to_tuple_list(test_dictionary)
```

_returned_ output:
```python
[
    ('Dr. No', 'Dr. No'),
    ('Mr. Big', 'Live and Let Die'),
]
```

```python
bond_villains = {
    'Dr. No': 'Dr. No',
    'Auric Goldfinger': 'Goldfinger',
    'Blofeld': 'You Only Live Twice',
    'Mr. Big': 'Live and Let Die'
}
```

<!-- grade_id: cell-1cc9f0a1f37968e0, points: 0, grade: false, solution: true -->
```python
# Complete the implementation of this function
def dict_to_tuple_list(d):
    return [(k, v) for k, v in d.items()]


```

<!-- grade_id: cell-24253ee970555d9b, points: 1, grade: true, solution: false -->
```python
# Question 5 Grading Checks

assert dict_to_tuple_list, "Make sure you name your function 'dict_to_tuple_list'."
```

#### Question 6

Write a function called `tuple_list_to_dict` that takes in **one list parameter** that's a list of tuples where each tuple contains the name of the villain and the movie they appeared in. (See the defined `bond_villans_list` in the next cell).

The function then **returns** a _dictionary_ where the keys are the names of the villains and the values are the movies they appeared in.

Example of function:
```python
test_list = [
    ('Dr. No', 'Dr. No'),
    ('Mr. Big', 'Live and Let Die')
]

tuple_list_to_dict(test_dictionary)
```

_returned_ output:
```python
{
    'Dr. No': 'Dr. No',
    'Mr. Big': 'Live and Let Die'
}
```


```python
bond_villains_list = [
    ('Dr. No', 'Dr. No'),
    ('Auric Goldfinger', 'Goldfinger'),
    ('Blofeld', 'You Only Live Twice'),
    ('Mr. Big', 'Live and Let Die')
]
```

<!-- grade_id: cell-d97414301dd87b76, points: 0, grade: false, solution: true -->
```python
# Complete the implementation of this function
def tuple_list_to_dict(l):
    return {k: v for k, v in l}

```

```python
tuple_list_to_dict(bond_villains_list)
```

Output:
```
{'Dr. No': 'Dr. No',
 'Auric Goldfinger': 'Goldfinger',
 'Blofeld': 'You Only Live Twice',
 'Mr. Big': 'Live and Let Die'}
```

<!-- grade_id: cell-c53b57b32208d368, points: 1, grade: true, solution: false -->
```python
# Question 6 Grading Checks

assert tuple_list_to_dict, "Make sure you name your function 'tuple_list_to_dict'."
```

```python

```